#Ressource produite dans le cadre du projet Mathéma-TIC (Samuel Bernard - 
#Cégep régional de Lanaudière à Terrebonne, et Bruno Poellhuber - Université de 
#Montréal), financé par le programme d’arrimage universités-collèges (PCUC) du 
#ministère de l'Enseignement supérieur, de la Recherche et de la Science du 
#Québec (MESRS) ainsi que par le Collège de Maisonneuve.

## DBsubject(Algèbre linéaire)
## DBchapter(Vecteurs et espace euclidien)
## KEYWORDS('R^3', 'produit vectoriel', 'vecteur', 'vecteur unitaire', 'vecteurs orthogonaux')
## Adaptation(Mélisande Fortin Boisvert - Collège de Maisonneuve)
## Traduction française : Sylvie Charbonneau
##########################################SOURCE################## 
## ZDate(5/31/2000)
## ZInstitution(Rochester)
## ZAuthor(Joseph Neisendorfer)
## ZTitleText1('Calculus: Early Transcendentals')
## ZAuthorText1('Stewart')
## ZEditionText1('6')
## ZSection1('12.4')
## ZProblem1('34')
## ZTitleText2('Calculus: Early Transcendentals')
## ZAuthorText2('Rogawski')
## ZEditionText2('1')
## ZSection2('12.4')
## ZProblem2('34')
##########################################SOURCE##################
##Level(2)
##Language(fr)

         
DOCUMENT();	# This should be the first executable line in the problem.

loadMacros("PGstandard.pl",
  "MathObjects.pl",
  "parserVectorUtils.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl");
 
TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Vector");
Context()->operators->undefine(".","><");

              

$a=Vector(1,random(2,10,2),1);
$b=Vector(1,random(3,10,2),1);
@a = $a->value; 
@b = $b->value; 
$rep1=unit($a x $b);
$rep2=-1*unit($a x $b);


$multians = MultiAnswer($rep1)->with(
  singleResult => 0,
  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my ($stu) = @{$student};
      if ( $stu  == $rep1 || $stu==$rep2) 
           {
            return [1];
           }
  }
);



Context()->texStrings;
BEGIN_TEXT
Sachant que \( \vec{a} = \vec{i} + $a[1]\vec{j} +  \vec{k} \) et 
\( \vec{b} = \vec{i} + $b[1] \vec{j} + \vec{k} \), 
déterminer un vecteur unitaire qui est orthogonal à la fois à 
\( \vec{a} \) et \( \vec{b} \).
$BR
Réponse: \{$multians->ans_rule(10)\}
$PAR
END_TEXT
Context()->normalStrings;
ANS( $multians->cmp() ); 

ENDDOCUMENT();		# This should be the last executable line in the problem.
