##DESCRIPTION
##  Multiplications et simplifications des fonctions rationnelles
##   avec restrictions
##ENDDESCRIPTION

## DBsubject(Algèbre)
## DBchapter(Fractions rationnelles)
## DBsection(Opérations sur les fractions rationnelles)
## KEYWORDS('simplification','fonction rationnelle','français','multiplication')
## Author(Julie Tremblay)
## Institution(Collège de Bois-de-Boulogne)

## Level(2)
## Language(fr)

DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "PGchoicemacros.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGcourse.pl",
);

TEXT(beginproblem());

######################################
#  Setup
#  Form: (x-a)(x+a)/(x+a)(x+b) * (x+b)/x, -a neq b

Context("Numeric");

$a = list_random(-5,-4,-3,-2,2,3,4,5);
do {$b = non_zero_random(-5,5,1);} while ($a==-1*$b);

@polynomials = (Formula("(x^2 - $a*$a)/(x^2 + ($a+$b)x + $a*$b)")->reduce->TeX, Formula("(x + $b)/x")->reduce->TeX);
@slice = NchooseK(2,2);

@f = @polynomials[@slice];

$ans1 = Compute("x-$a");
$ans2 = Compute("x");


$mp = MultiAnswer($ans1,$ans2)->with(
      singleResult => 0,
      checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
			  Value::Error("As-tu simplifié ta réponse?") 
                if ($quotientValid && !$numeratorValid && !$denominatorValid);
			Value::Error("Ton numérateur et dénominateur sont incorrects.") 
                unless ($numeratorValid || $denominatorValid);
			  Value::Error("Ton numérateur est incorrect") 
                unless $numeratorValid;
              Value::Error("Ton dénominateur est incorrect") 
                unless $denominatorValid;
              Value::Error("Une de tes réponses est de signe contraire") 
                unless $quotientValid;
            }

            
		 if ($quotientValid) {
			return[$numeratorValid,$denominatorValid];
         } else {
            return[0,0];
        }
    });


######################################
#  Main text

BEGIN_TEXT
Effectuer la multiplication suivante et simplifier la réponse en donnant les restrictions sur le domaine.$BR
$BR
\[$f[0] \cdot $f[1]\]
$PAR
Exprimer la réponse en fraction rationnelle simplifiée \(\dfrac{A}{B}\). $PAR 
$BR
Réponse : \(A =\) \{$mp->ans_rule(20)\} et \(B =\) \{$mp->ans_rule(20)\}
$PAR
Restriction(s) : \(x \neq\) \{ ans_rule(10) \}
$BR $BR
${BITALIC}${BBOLD}Note${EBOLD} : S'il y a plusieurs restrictions de \(x\), les séparer par une virgule. (Exemple : \(x \neq 1, –4,2\))$EITALIC
END_TEXT


######################################
#  Answer

$showPartialCorrectAnswers = 1;
ANS( $mp->cmp );


if($a==$b){
$rest = List("0,-$a");
} else{
$rest = List("0,-$a,-$b");
}
ANS($rest->cmp);



BEGIN_SOLUTION
$PAR 
En factorisant et en effectuant la multiplication, on obtient :
$PAR
\[$f[0] \cdot $f[1]=\frac{(x-$a)(x+$a)(x+$b)}{(x+$a)(x+$b)x}\]
$PAR
Ainsi, les restrictions sur le domaine sont \(x \not\in \lbrace $rest\rbrace \), car ce sont les valeurs de \(x\) pour lesquelles les dénominateurs sont égaux à 0. En simplifiant les facteurs communs au numérateur et au dénominateur, la réponse devient
$PAR
\[\begin{array}{ll} \displaystyle $f[0] \cdot $f[1]&= \dfrac{(x-$a)\cancel{(x+$a)}\cancel{(x+$b)}}{\cancel{(x+$a)}\cancel{(x+$b)}x}\\&= \dfrac{(x-$a)}{x}  \text{ si }  x \neq \lbrace $rest \rbrace \end{array}\]

END_SOLUTION


######################################


COMMENT('MathObject version');
ENDDOCUMENT();
