##DESCRIPTION
##  Algebra: Summing rational expressions 
##ENDDESCRIPTION

## DBsubject(Fonctions et manipulations algébriques)
## DBchapter(Opérations algébriques et factorisation)
## KEYWORDS('addition', 'fraction rationnelle', 'simplification')
## Language(fr)
## Traduction('Julie Tremblay')
## Institution(Collège de Bois-de-Boulogne)



########################################################################

DOCUMENT();      

loadMacros(
   "PGstandard.pl",     # Standard macros for PG language
   "MathObjects.pl",
   "parserMultiAnswer.pl",
   "PGcourse.pl",      # Customization file for the course
);

########################################################################

TEXT(beginproblem());

$showPartialCorrectAnswers = 1;

########################################################################

Context("Numeric");
Context()->variables->are(n=>'Real'); 

$a = random(2,10,1);
$b = random(2,10,2);

do{$c = random(4,8,1);} until(gcd($c,3)==1);

do{$d = random(2,10,1);} until(gcd($b,$d)==1 && gcd($c,$d)==1);


$poly1 = Compute("3*$a") -> reduce;
$poly2 = Compute("$a*$c*n^2") -> reduce;
$poly3 = Compute("$b") -> reduce;
$poly4 = Compute("$d*n") -> reduce;

$co1 = 3*$d;
$co2 = $b*$c;
$co3 = $c*$d;

$ans1 = Compute("$co1+$co2 n")->reduce;
$ans2 = Compute("$co3 n^2")->reduce;

########################################################################

$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
		 allowBlankAnswers => 1,
         checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom

            my $isPreview = $main::inputs_ref->{previewAnswers};

            
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

            if (! $isPreview) { # Don't show messages if previewing....
              Value::Error("Votre numérateur est incorrect") 
                unless $numeratorValid;
              Value::Error("Votre dénominateur est incorrect") 
                unless $denominatorValid;
              Value::Error("Une de vos réponses est de signe contraire") 
                unless $quotientValid;
            }

            return $quotientValid;
            },
        );

########################################################################

Context()->texStrings;
BEGIN_TEXT
Effectuer l'addition suivante.
\[
\frac{$poly1}{$poly2} + \frac{$poly3}{$poly4}
\]
$BR
Exprimer la réponse sous la forme simplifiée  \(\dfrac{A}{B}\).$BR $BR
Réponse : \(A =\) \{$mp->ans_rule(20)\} et \(B =\) \{$mp->ans_rule(20)\}
END_TEXT
Context()->normalStrings;

########################################################################

ANS( $mp->cmp );

ENDDOCUMENT();