##DESCRIPTION
##  Rational Expressions and Functions: Multipy and Divide
## 
##ENDDESCRIPTION

## DBsubject(Fonctions et manipulations algébriques)
## DBchapter(Opérations algébriques et factorisation)
## KEYWORDS('addition', 'domaine', 'fraction rationnelle', 'restriction', 'simplification')
## Author(Julie Tremblay)
## Institution(Collège de Bois-de-Boulogne)

## Level(2)
## Language(fr)



DOCUMENT(); # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
 "PGcourse.pl",
  "MathObjects.pl",
 "parserMultiAnswer.pl",
);



TEXT(beginproblem());

######################################
#  Setup
#  Forme: A/x + -Ab/(x-a)^2  - 2x/(x-a)(x+b)

Context("Numeric");
Context()->noreduce('(-x)-y','(-x)+y');

$a= random(1,5,1);
$b= random(1,5,1);

if($a==$b){$b++;}

$A = non_zero_random(-3,3,1);



#les coefficients du numérateur
$c0 = Compute("$A*$a*$a*$b")->reduce; 
$c1 = Compute("-2*$a*$b*$A + $A*$a*$a-$A*$b*$b")->reduce;
$c3 = Compute("-$A");

$p1 = Compute("$b-$a");
$p2 = Compute("$a*$b");
$p3 = Compute("-2*$A");

$poly1= Compute("x^2+$p1*x-$p2")->reduce;
$poly2 = Compute("-$A*$b")->reduce;

$f = Compute("$A/x+ $poly2/(x-$a)^2 + $p3 x/$poly1");

$ans1 = Formula("$c3*x^3 + $c1*x + $c0")->reduce;
$ans2 = Formula("x*(x-$a)^2*(x+$b)")->reduce;



$mp = MultiAnswer($ans1,$ans2)->with(
         singleResult => 0,
        checker => sub {
            my ($correct,$student,$self) = @_;  # get the parameters
            my ( $p1, $p2 ) = @{$correct};      # get the correct num, denom
            my ( $p1stu, $p2stu ) = @{$student};# get the student's num, denom
          
            $numeratorValid   = ( $p1==$p1stu || Compute("-1")*$p1==$p1stu );
            $denominatorValid = ( $p2==$p2stu || Compute("-1")*$p2==$p2stu );
            $quotientValid    = ( $p1/$p2 == Compute($p1stu)/Compute($p2stu) );

	
			if ($quotientValid) {
				return[$numeratorValid,$denominatorValid];
			} elsif ($numeratorValid && $denominatorValid ) {
				return[0,$denominatorValid];
			} elsif  ($numeratorValid || $denominatorValid ) {
			   return[$numeratorValid,$denominatorValid];
			} else {
                           return[0,0];
                        }
		}
        );
######################################
#  Main text


Context()->texStrings;
BEGIN_TEXT
Effectuer l'opération suivante.
\[$f\]
$BR
$PAR
Exprimer la réponse en fraction rationnelle simplifiée \(\dfrac{A}{B}\) 
et donner la ou les restrictions sur le domaine de l'expression. 
$PAR 
Réponse : \(A =\) \{$mp->ans_rule(20)\} et \(B =\) \{$mp->ans_rule(20)\} 
$PAR 
Restriction(s) : \(x\neq\) \{ans_rule(30)\}
$BR $BR
${BITALIC}${BBOLD}Note${EBOLD} : S'il y a plusieurs restrictions de \(x\), les séparer par une virgule. (Exemple : \(x \neq 1, –4\))$EITALIC
END_TEXT
Context()->normalStrings;

######################################
#  Answer

ANS( $mp->cmp );

$rest3 = Compute("-$b"); 

ANS(List(0,$a,$rest3)->cmp(
list_type  => 'liste',
));

$showPartialCorrectAnswers = 1;

######################################

$ans = Formula("$answer")->reduce->TeX;


$poly3 = Compute("x^2-2*$a*x+$a^2")->reduce;
$poly4 = Compute("x^3+$b*x^2 - 2*$a*x^2 - 2*$a*$b*x+($a)^2 x + ($a)^2*$b")->reduce;
$q2 = Compute("$A*$b - $A*2*$a")->reduce;
$q1 = Compute("-2*$A*$a*$b + $A*($a)^2")->reduce;
$r1 = Compute("$poly2*$b");
$r2 = Compute("-$p3*$a");


Context()->texStrings;
BEGIN_SOLUTION
$PAR
En factorisant le troisième dénominateur, on obtient \($poly1 = (x-$a)(x+$b)\).
$BR
Le plus petit dénominateur commun aux trois fractions rationnelles est : 
$BR
\[  $ans2\]
$BR
Ainsi les restrictions sur le domaine sont : \(x\neq 0, $a\) et \($rest3\). $BR
Voici la démarche pour effectuer l'opération sur les fractions :
$PAR
\[\begin{array}{l}  \displaystyle $f \\
= \dfrac{$A}{x} + \dfrac{$poly2}{(x-$a)^2} + \dfrac{$p3 x}{(x-$a)(x+$b)}\\
= \dfrac{$A(x-$a)^2(x+$b)}{$ans2}+\dfrac{$poly2 x(x+$b)}{$ans2}+\dfrac{$p3 x\cdot x(x-$a)}{$ans2}\\
= \dfrac{$A($poly3)(x+$b)+$poly2 x(x+$b) + $p3 x^2(x-$a)}{$ans2}\\
= \dfrac{$A($poly4) + $poly2 x^2 +$r1 x +$p3 x^3 + $r2 x^2}{$ans2} \\
=  \dfrac{$A x^3 + \cancel{$q2 x^2} + $q1 x + $c0 + \cancel{$poly2 x^2} +$r1 x + $p3 x^3 +\cancel{$r2 x^2}}{$ans2} \\
= \dfrac{$ans1}{$ans2}\end{array}
\]
END_SOLUTION
Context()->normalStrings;

COMMENT('MathObject version');
ENDDOCUMENT();