## DBsubject(Calcul différentiel et intégral à plusieurs variables)
## DBchapter(Applications de la limite et de la dérivée)
## KEYWORDS('applications', 'courbes de niveau', 'méthode du gradient')
## Modifié par Éric Gaul, Collège Montmorency
## Traduction française : Éric Gaul
## Date : mars 2017
## Language(fr)


###############Source###############
# DESCRIPTION
# Problem from Calculus, multi-variable, Hughes-Hallett et al.,
# originally from 5ed (with updates)
# WeBWorK problem written by Gavin LaRose, <glarose@umich.edu>
# ENDDESCRIPTION
## Tagged by glr 12/15/09
## ZDBsubject(Calculus - multivariable)
## ZDBchapter(Differentiation of multivariable functions)
## ZDBsection(Directional derivatives and the gradient)
## Institution(University of Michigan)
## Author(Gavin LaRose)
## Level(4)
## MO(1)
## TitleText1('Calculus')
## AuthorText1('Hughes-Hallett')
## EditionText1('5')
## Section1('14.5')
## Problem1('51')
## TitleText2('Calculus')
## AuthorText2('Hughes-Hallett')
## EditionText2('6')
## Section2('14.5')
## Problem2('52')
## Textbook tags
## HHChapter1('Differentiating Functions of Several Variables')

DOCUMENT();

loadMacros(
"PGstandard.pl",
"PGchoicemacros.pl",
"MathObjects.pl",
);

Context("Vector2D");
Context()->flags->set( ijk=>1 );
Context()->constants->set(
  i => {TeX => "\,\mathit{\vec i}"},
  j => {TeX => "\,\mathit{\vec j}"},
);
$showPartialCorrectAnswers = 1;

$dummy = random(1,5,1);
$a = random(1,5,1);
$trig = ("cos", "sin")[random(0,1,1)];
$trig2 = ($trig eq "cos") ? "-sin" : "cos";

$zfunc = Compute( "(e^x - $a*x)*$trig(y)" )->reduce();

$x0 = random(1,5,1);
$y0 = random(2,6,1);

$fx0 = Compute( "(e^($x0) - $a)*$trig($y0)" );
$fy0 = Compute( "(e^($x0) - $a*$x0)*($trig2($y0))" );
$grad0 = Vector( "-$fx0 i - $fy0 j" );

$vx0 = random(1,8,1);
$vy0 = random(1,8,1);
$aval = Compute( "$fx0*$vx0 + $fy0*$vy0" );

Context()->texStrings;
TEXT(beginproblem());
BEGIN_TEXT
Soit la surface \( S \) définie par \( z=f(x,y) \), où
\( f(x,y) = $zfunc \).
$PAR
${BBOLD}(a)$EBOLD
Trouver un vecteur perpendiculaire à la courbe de niveau de
\( f \) passant par le point \( ($x0,$y0) \), dans la direction
où \( f \) décroît le plus rapidement.
$PAR
\{ ans_rule(55) \} $BITALIC (ex. indiquer un vecteur comme \((0, 1)\) par : \( < 0, 1 > \)) $EITALIC

$PAR
${BBOLD}(b)$EBOLD
Supposons que \( \vec v = $vx0\vec i + $vy0 \vec j + a \vec k \)
est un vecteur de \( \mathbb{R}^3 \), tangent à la surface \( S \) au point \( P($x0, $y0, f($x0, $y0) ) \).
Quelle est la valeur de \( a \) ?
$PAR
\(\ \ \ a = \) \{ ans_rule(35) \}

END_TEXT
Context()->normalStrings;

ANS($grad0->cmp( checker=>sub {
    my ($correct, $student, $ansHash) = @_;
    return $correct->isParallel($student);
  }
) );
# ANS($grad0->cmp() );
ANS($aval->cmp() );

$trigtex = '\' . $trig;
if ( $trig2 eq 'cos' ) {
    $trig2tex = '\cos';
    $trig2sgn = '+';
    $ntrig2sgn = '-';
} else {
    $trig2tex = '\sin';
    $trig2sgn = '-';
    $ntrig2sgn = '+';
}
$trig2sgnd = $trig2sgn eq '+' ? '' : '-';
$ad = $a == 1 ? '' : $a;
$ax0 = $a*$x0;

$ex0 = ( $x0 == 1 ) ? 'e' : "e^{$x0}";

Context()->texStrings;
SOLUTION(EV3(<<'END_SOLUTION'));
${BBOLD}(a)$EBOLD
Le vecteur gradient \(\nabla f(x,y)\) est perpendiculaire à la courbe
de niveau de \(f\) qui passe par \((x,y)\) :
\[
\nabla f(x,y) = (e^x-$a)$trigtex y\vec i $trig2sgn (e^x- $ad x)$trig2tex y\vec j.
\]
Donc, au point \(($x0,$y0)\) :
\[
\nabla f($x0,$y0)=($ex0-$a)$trigtex $y0\vec i $trig2sgn
    ($ex0-$ax0)$trig2tex $y0\vec j
\]
Le vecteur \(\nabla f\) pointe dans la direction de la variation de
\(f\) la plus élevée, donc le vecteur cherché peut être :
\[
-\nabla f($x0,$y0) = -($ex0-$a)$trigtex $y0\vec i $ntrig2sgn
    ($ex0-$ax0)$trig2tex $y0\vec j.
\]

$PAR
${BBOLD}(b)$EBOLD
Afin de trouver un vecteur normal à la surface, on écrit
cette dernière sous la forme :
\[
F(x,y,z)=(e^x- $ad x)$trigtex y-z=0.
\]
Donc :
\[
\nabla F = (e^x - $a)$trigtex y\vec i $trig2sgn (e^x - $ad x)$trig2tex y\vec j
   - \vec k.
\]
Ainsi, en \(P\) :
\[
\nabla F = ($ex0 - $a)$trigtex $y0\vec i $trig2sgn
    ($ex0 - $ax0)$trig2tex y\vec j - \vec k.
\]
Le vecteur \(\vec v\) doit être perpendiculaire à \(\nabla F\), donc :
\[
\vec v\cdot\nabla F =
($vx0\vec i + $vy0\vec j + a\vec k)\cdot \nabla F($x0,$y0) =0.
\]
En résolvant pour \( a \) : 
\[ a = $vx0 (($ex0-$a)$trigtex $y0) + $vy0 ($trig2sgnd
    ($ex0-$ax0)$trig2tex $y0). \]
END_SOLUTION
Context()->normalStrings;


;
ENDDOCUMENT();
