## DBsubject(Algèbre linéaire)
## DBchapter(Espace vectoriel)
## KEYWORDS('base', 'ensemble générateur', 'espace euclidien', 'indépendance')
## Auteure(Mélisande Fortin Boisvert - Collège de Maisonneuve)
##Level(3)
##Language(fr)

DOCUMENT();

loadMacros(
"PGstandard.pl",
"MathObjects.pl",
"parserMultiAnswer.pl",
"PGcourse.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Vector');


$rep1= Vector(1,0,0);
$rep2= Vector(0,1,0);
$rep3= Vector(0,0,1);
   


#Un bonne réponse
#C contiendra cette bonne réponse
#S contiendra la réponse de l'étudiant


$image_multians = MultiAnswer($rep1, $rep2, $rep3)->with(
  singleResult => 1,
  separator => ',',
  tex_separator => ',',

  checker => sub {
      my ( $correct, $student, $self ) = @_;
      my @c = @{$correct};
      my @s = @{$student};

      return 0 if scalar(@s) < scalar(@c);  # Valide s'il manque des entrées

     
	  $C=Matrix($rep1 ,$rep2, $rep3
          )->transpose;

    
  
      $S = Matrix( $s[0] , $s[1], $s[2]
      )->transpose;

	  
	  #Les réponses de l'étudiant sont les colonnes de la matrice $S
      # A^T A est inversible ssi les colonnes sont linéairement idépendantes


	  
	#Validation: les vecteurs ne sont pas parallèles  à i  et j et k
	  #  Check that the student's vectors are linearly independent
      if ( $rep1->isParallel($s[0]) || $rep1->isParallel($s[1])  || $rep1->isParallel($s[2]) 
			||$rep2->isParallel($s[0])  || $rep2->isParallel($s[1]) || $rep2->isParallel($s[2])
				||$rep3->isParallel($s[0])  || $rep3->isParallel($s[1]) || $rep3->isParallel($s[2])){
        Value->Error("Las vecteurs ne doivent pas être parallèles à \( \vec{i}\) ,  \( \vec{j} \) ou\( \vec{k} \)  ");
        return 0;
      }
     

      #Validation: les vecteurs sont linéairement indépendants
	  
      if ( (($S->transpose) * $S)->det == 0) {
        Value->Error("Les vecteurs sont linéairement dépendants");
        return 0;
      }

	  
	   #Validation: les vecteurs sont orthogonaux
      if ( $s[0].$s[1] != 0 || $s[0].$s[2] != 0 || $s[1].$s[2] != 0 ) {
        Value->Error("Les vecteurs ne sont pas orthogonaux deux à deux");
        return 0;
      }
	  
	  
	  
     	  #Validation: les vecteurs engendrent le bon espace vectoriel
		  #  C et S engendrent le même espace s'il existe un changement de base X 
		  #  S = CX et  X est donné par (C^T C)^{-1} C^T S = X.
	      $CTC = ($C->transpose) * $C;
          $X = ($CTC->inverse) * (($C->transpose) * $S);
      return $S == $C * $X;


  }
);



Context()->texStrings;
BEGIN_TEXT

Trouver une base \(B\)  orthogonale de \( \mathbb{R}^3   \) dont les vecteurs ne sont pas parallèles à \(\vec{i}\),  \(\vec{j}\) et \(\vec{k}\).
$BR
$BR
\(B\)  = \(\Bigg\lbrace\) 
\{ $image_multians->ans_array() \}, 
\{ $image_multians->ans_array() \},
\{ $image_multians->ans_array() \},
\(\Bigg\rbrace\)
END_TEXT
Context()->normalStrings;

ANS( $image_multians->cmp() );


;

ENDDOCUMENT();
